# Đồ án cuối kỳ - Môn Lập trình khoa học dữ liệu
# Giáo viên giảng dạy: Thầy Trần Trung Kiên
## Sinh viên thực hiện:

Tạ Võ Anh Khuê - 19120551

Nguyễn Minh Long - 19120568

---

## Import các thư viện cần thiết để sử dụng

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd 
```

---

## Thu thập dữ liệu cho đồ án - Currently Working

Dữ liệu được sử dụng trong bài tập này là dữ liệu về các phàn nàn, khiếu nại của người dùng được gửi tới CFBP(Consumer Financial Protection Bureau) - tạm gọi là Cục Bảo Vệ Tài Chính Cho Người Tiêu Dùng. Dữ liệu được lấy [ở đây](https://www.kaggle.com/meetnagadia/consumer-complaint-finance). 

Cục Bảo Vệ Tài Chính Cho Người Tiêu Dùng (CFPB) là một cơ quan của chính phủ liên bang Mỹ được thành lập để bảo đảm sự hiệu quả của những sản phẩm và dịch vụ tài chính cho mọi người như mua nhà, chọn lựa một thẻ tín dụng, gửi tiền cho người thân ở nước ngoài, hoặc khi sử dụng bất cứ loại sản phẩm tài chánh khác ở Mỹ.
Cục có chức năng giám sát hoạt động của các ngân hàng, liên hiệp tín dụng (credit union), những công ty tài chính khác và thi hành những quy luật tài chính để bảo vệ người dùng tránh khỏi các hành vi bất công, dối trá và lạm dụng.

Đôi nét về dữ liệu này: Dữ liệu thể hiện các complaint(khiếu nại, phàn nàn) của người dùng về các vấn đề tài chính mà họ gặp phải. Người dùng trình báo lên Cục Bảo Vệ(CFBP) và cục có trách nhiệm liên lạc và trao đổi, chuyển tiếp thông tin giúp người dùng tới các công ty liên quan. Dữ liệu được công khai ghi nhận(published) bởi CFPB sau khi mà các công ty phản hồi, xác nhận có liên quan đến các vấn đề của người dùng hoặc sau 15 ngày, tùy điều kiện nào đến trước.

---

## Khám phá dữ liệu - Currently Working

### Đọc dữ liệu từ file

Dữ liệu được đặt chung cấp folder của file noteboook, đọc vào bằng hàm read_csv đơn giản, không có gì phức tạp trong phần này.

<!-- grade_id: cell-d46e97079096473c, points: 0.25, grade: true, solution: true -->
```python
df = pd.read_csv('complaints.csv')
```

Output:
```
D:\Miniconda\envs\min_ds-env\lib\site-packages\IPython\core\interactiveshell.py:3172: DtypeWarning: Columns (9) have mixed types.Specify dtype option on import or set low_memory=False.
  has_raised = await self.run_ast_nodes(code_ast.body, cell_name,
```

```python
df.head()
```

Output:
```
Date received                                            Product  \
0    2019-06-13  Credit reporting, credit repair services, or o...   
1    2019-11-01                              Vehicle loan or lease   
2    2019-04-01  Credit reporting, credit repair services, or o...   
3    2019-04-10                        Credit card or prepaid card   
4    2020-04-21  Credit reporting, credit repair services, or o...   

         Sub-product                                 Issue  \
0   Credit reporting  Incorrect information on your report   
1               Loan           Struggling to pay your loan   
2   Credit reporting  Incorrect information on your report   
3  Store credit card                  Closing your account   
4   Credit reporting  Incorrect information on your report   

                             Sub-issue  \
0  Information belongs to someone else   
1     Denied request to lower payments   
2             Account status incorrect   
3             Can't close your account  
```

```python
df['Consumer disputed?'].unique()
```

Output:
```
array([nan, 'No', 'Yes'], dtype=object)
```

### Mỗi cột dữ liệu có ý nghĩa là gì?

Data có tổng cộng 18 cột, với các ý nghĩa cơ bản của từng cột như sau:
* Date received: ngày nhận được khiếu nại
* Product: loại sản phẩm, dịch vụ khiếu nại,....
* Sub product: thêm thông tin, làm rõ hơn về sản phẩm, dịch vụ,... trong product
* Issue: vấn đề người dùng gặp phải
* Sub-issue: thông tin thêm, làm rõ về vấn đề gặp phải trong issue
* Consumer complaint narrative: miêu tả, phàn nàn, khiếu nại của người dùng đối với vấn đề họ gặp phải
* Company public response: phản hồi chính thức, công khai của phía công ty
* Company: tên công ty
* State: tên bang
* ZIP code : mã zip ( giống mã bưu điện )
* Consumer consent provided: 
* Submitted via: khiếu nại được gửi thông qua dạng nào(web,fax,email,...)
* Date sent to company: ngày khiếu nại được gửi đến công ty
* Timely response:
* Consumer disputed?: người khiếu nại có mở tranh chấp không(dạng đòi bồi thường này kia)
* Complaint ID: id của đơn khiếu nại

### Dữ liệu có bao nhiêu dòng và bao nhiêu cột? (0.25đ)

Kế đến, bạn tính số dòng và số cột của DataFrame `survey_df` và lần lượt lưu vào biến `num_rows` và `num_cols`. 

<!-- grade_id: cell-8338f3610f6c5002, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
num_rows = len(survey_df)
num_cols = len(survey_df.columns)
#raise NotImplementedError()
```

### Mỗi dòng có ý nghĩa gì? Có vấn đề các dòng có ý nghĩa khác nhau không?

Theo file "README_2020.txt" cũng như theo quan sát sơ bộ về dữ liệu, mỗi dòng trong DataFrame `survey_df` cho biết kết quả làm khảo sát của một người. Có vẻ không có vấn đề các dòng có ý nghĩa khác nhau. 

### Dữ liệu có các dòng bị lặp không? (0.25đ)

Kế đến, bạn tính số dòng có index (id của người làm khảo sát) bị lặp và lưu vào biến `num_duplicated_rows`. Trong nhóm các dòng có index giống nhau thì dòng đầu tiên không tính là bị lặp.

<!-- grade_id: cell-aadbfe12be205faa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
num_duplicated_rows = 0
check_list = survey_df.index.duplicated(keep ='first')
for value in check_list:
    if value == True:
        num_duplicated_rows += 1
#raise NotImplementedError()
```

### Mỗi cột có ý nghĩa gì? (0.25đ)

Để xem ý nghĩa của mỗi cột thì:
- Trước tiên, bạn cần đọc file "survey_results_schema-short.csv" vào DataFrame `col_meaning_df`; bạn cũng cần cho cột "Column" làm cột index. 
- Sau đó, bạn chỉ cần hiển thị DataFrame `col_meaning_df` ra để xem (vụ này khó nên ở dưới mình đã làm cho bạn ở cell có dòng "# TEST" 😉). Tuy nhiên, bạn sẽ thấy ở cột "QuestionText": các chuỗi mô tả bị cắt do quá dài. Do đó, trước khi hiển thị DataFrame `col_meaning_df`, bạn cũng cần chỉnh sao đó để các chuỗi mô tả không bị cắt (vụ này bạn tự search Google, gợi ý: bạn sẽ dùng đến câu lệnh `pd.set_option`).   

<!-- grade_id: cell-5efb0517df207535, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
col_meaning_df = pd.read_csv('survey_results_schema-short.csv',index_col=0)
pd.set_option('display.max_columns', None)
pd.set_option('display.expand_frame_repr', False)
pd.set_option('max_colwidth', 800)
#raise NotImplementedError()
```

```python
# TEST
col_meaning_df
```

Trước khi đi tiếp, bạn nên đọc kết quả hiển thị ở trên và đảm bảo là bạn đã hiểu ý nghĩa của các cột. Để hiểu ý nghĩa của cột, có thể bạn sẽ cần xem thêm các giá trị của cột bên DataFrame `survey_df`. 

### Mỗi cột hiện đang có kiểu dữ liệu gì? Có cột nào có kiểu dữ liệu chưa phù hợp để có thể xử lý tiếp không? (0.25đ)

Kế đến, bạn tính kiểu dữ liệu (dtype) của mỗi cột trong DataFrame `survey_df` và lưu kết quả vào Series `dtypes` (Series này có index là tên cột). 

<!-- grade_id: cell-206c19691b1a6c05, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
dtypes = survey_df.dtypes
#raise NotImplementedError()
dtypes
```

<!-- grade_id: cell-3ebcf8f07733605b, points: 0.25, grade: true, solution: false -->
```python
# TEST
float_cols = set(dtypes[(dtypes==np.float32) | (dtypes==np.float64)].index)
assert float_cols == {'Age', 'ConvertedComp', 'WorkWeekHrs'}
object_cols = set(dtypes[dtypes == object].index)
assert len(object_cols) == 25
```

Như bạn có thể thấy, cột "YearsCode" và "YearsCodePro" nên có kiểu dữ liệu số, nhưng hiện giờ đang có kiểu dữ liệu object. Ta hãy thử xem thêm về các giá trị 2 cột này. 

```python
survey_df['YearsCode'].unique()
```

```python
survey_df['YearsCodePro'].unique()
```

Ta nên đưa 2 cột này về dạng số để có thể tiếp tục khám phá (tính min, median, max, ...).

---

## Tiền xử lý (0.5đ)

Bạn sẽ thực hiện tiền xử lý để chuyển 2 cột "YearsCode" và "YearsCodePro" về dạng số (float). Trong đó: "Less than 1 year" $\to$ 0, "More than 50 years" $\to$ 51. Sau khi chuyển thì `survey_df.dtypes` sẽ thay đổi.

<!-- grade_id: cell-cfc7d3af0fb5edec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
survey_df = survey_df.replace(to_replace = 'Less than 1 year',value= '0')
survey_df = survey_df.replace(to_replace = 'More than 50 years',value= '51')
survey_df['YearsCode'] = pd.to_numeric(survey_df['YearsCode'])
survey_df['YearsCodePro'] = pd.to_numeric(survey_df['YearsCodePro'])
#raise NotImplementedError()
survey_df.dtypes
```

<!-- grade_id: cell-4ee41f84604837a1, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert survey_df['YearsCode'].dtype in [np.float32, np.float64]
assert survey_df['YearsCodePro'].dtype in [np.float32, np.float64]
```

---

## Quay lại bước khám phá dữ liệu 

### Với mỗi cột có kiểu dữ liệu dạng số, các giá trị được phân bố như thế nào? (1đ)

(Trong đó: phần tính các mô tả của mỗi cột chiếm 0.5đ, phần tính số lượng giá trị không hợp lệ của mỗi cột chiếm 0.5đ)

Với các cột có kiểu dữ liệu số, bạn sẽ tính:
- Tỉ lệ % (từ 0 đến 100) các giá trị thiếu 
- Giá trị min
- Giá trị lower quartile (phân vị 25)
- Giá trị median (phân vị 50)
- Giá trị upper quartile (phân vị 75)
- Giá trị max

Bạn sẽ lưu kết quả vào DataFrame `nume_col_info_df`, trong đó: 
- Tên của các cột là tên của các cột số trong `survey_df`
- Tên của các dòng là: "missing_percentage", "min", "lower_quartile", "median", "upper_quartile", "max"  

Để dễ nhìn, tất cả các giá trị bạn đều làm tròn với 1 chữ số thập phân bằng phương thức `.round(1)`.

<!-- grade_id: cell-d1503f01712c873a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
list_row = ['missing_percentage', 'min', 'lower_quartile', 'median', 'upper_quartile', 'max']
list_col_numeric = ['Age', 'ConvertedComp', 'WorkWeekHrs', 'YearsCode', 'YearsCodePro']
list_data = []
for i in range(0,5):
    list_data.append([]) # chứ thông tin mỗi thuộc tính
total_values = float(len(survey_df))
i = -1
for value in list_col_numeric: #xử lý cho từng thuộc tính
    i += 1
    describe = survey_df[value].describe()
    miss_value = total_values - describe[0]
    percent_miss = round((miss_value/total_values)*100, 1)
    list_data[i].append(percent_miss)
    list_data[i].append(describe[3].round(1))
    list_data[i].append(describe[4].round(1))
    list_data[i].append(describe[5].round(1))
    list_data[i].append(describe[6].round(1))
    list_data[i].append(describe[7].round(1))
data_numeric = {list_col_numeric[0]:list_data[0], list_col_numeric[1]:list_data[1],list_col_numeric[2]:list_data[2],list_col_numeric[3]:list_data[3], list_col_numeric[4]:list_data[4]}
nume_col_info_df = pd.DataFrame(data_numeric, index = list_row)
#raise NotImplementedError()
nume_col_info_df
```

<!-- grade_id: cell-b19aff83344102f8, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert nume_col_info_df.shape == (6, 5)
data = nume_col_info_df.loc[['missing_percentage', 'min', 'lower_quartile', 'median', 'upper_quartile', 'max'],
                           ['Age', 'ConvertedComp', 'WorkWeekHrs', 'YearsCode', 'YearsCodePro']].values
correct_data = np.array([[     29.5,      46.1,      36.2,      10.5,      28.1],
                         [      1. ,       0. ,       1. ,       0. ,       0. ],
                         [     24. ,   24648. ,      40. ,       6. ,       3. ],
                         [     29. ,   54049. ,      40. ,      10. ,       6. ],
                         [     35. ,   95000. ,      44. ,      17. ,      12. ],
                         [    279. , 2000000. ,     475. ,      51. ,      51. ]])
assert np.array_equal(data, correct_data)
```

**Có giá trị không hợp lệ trong mỗi cột không? (không xét giá trị thiếu)**
- Cột "Age": bạn hãy tính số lượng giá trị không hợp lệ của cột "Age" (< giá trị tương ứng trong cột "YearsCode" HOẶC < giá trị tương ứng trong cột "YearsCodePro") và lưu kết quả vào biến `num_invalid_Age_vals`.
- Cột "WorkWeekHrs" (số giờ làm việc trung bình một tuần): ta thấy max là 475 giờ! Trong khi đó, 7 ngày * 24 giờ = 168 giờ! Bạn hãy tính số lượng giá trị không hợp lệ của cột "WorkWeekHrs" (> 24 * 7) và lưu kết quả vào biến `num_invalid_WorkWeekHrs`.
- Cột "YearsCode": bạn hãy tính số lượng giá trị không hợp lệ của cột "YearsCode" (< giá trị tương ứng trong cột "YearsCodePro" HOẶC > giá trị tương ứng trong cột "Age") và lưu kết quả vào biến `num_invalid_YearsCode`.
- Cột "YearsCodePro": bạn hãy tính số lượng giá trị không hợp lệ của cột "YearsCodePro" (> giá trị tương ứng trong cột "YearsCode" HOẶC > giá trị tương ứng trong cột "Age") và lưu kết quả vào biến `num_invalid_YearsCodePro`.

<!-- grade_id: cell-59a58bdc5f54bb25, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
num_invalid_Age_vals = 0
num_invalid_WorkWeekHrs = 0
num_invalid_YearsCode = 0
num_invalid_YearsCodePro = 0
num_invalid_Age_vals = len(survey_df.loc[(survey_df['Age'] < survey_df['YearsCode']) | (survey_df['Age'] < survey_df['YearsCodePro'])])
num_invalid_WorkWeekHrs_vals = len(survey_df.loc[survey_df['WorkWeekHrs'] > 168])
num_invalid_YearsCode_vals = len(survey_df.loc[(survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCode'] > survey_df['Age'])])
num_invalid_YearsCodePro_vals = len(survey_df.loc[(survey_df['YearsCodePro'] > survey_df['YearsCode']) | (survey_df['YearsCodePro'] > survey_df['Age'])])
#raise NotImplementedError()

```

<!-- grade_id: cell-386769471c5b4b8f, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert num_invalid_WorkWeekHrs_vals == 62
assert num_invalid_Age_vals == 16
assert num_invalid_YearsCode_vals == 499
assert num_invalid_YearsCodePro_vals == 486
```

Do số lượng các giá trị không hợp lệ cũng khá ít nên ta có thể tiền xử lý bằng cách xóa các dòng chứa các giá trị không hợp lệ.

---

## Tiền xử lý (0.5đ)

Bạn sẽ thực hiện tiền xử lý để xóa đi các dòng của DataFrame `survey_df` mà chứa ít nhất là một giá trị không hợp lệ. Sau khi tiền xử lý thì `survey_df` sẽ thay đổi. 

<!-- grade_id: cell-6a6242ef6d2c8c65, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#chọn ra những dòng không hợp lệ rồi bỏ đi trong dataframe gốc
temp_data_frame = survey_df.loc[(survey_df['Age'] < survey_df['YearsCode']) | (survey_df['Age'] < survey_df['YearsCodePro'])]
survey_df.drop(temp_data_frame.index, axis=0,inplace=True)
temp_data_frame = survey_df.loc[survey_df['WorkWeekHrs'] > 168]
survey_df.drop(temp_data_frame.index, axis=0,inplace=True)
temp_data_frame = survey_df.loc[(survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCode'] > survey_df['Age'])]
survey_df.drop(temp_data_frame.index, axis=0,inplace=True)
temp_data_frame = survey_df.loc[(survey_df['YearsCodePro'] > survey_df['YearsCode']) | (survey_df['YearsCodePro'] > survey_df['Age'])]
survey_df.drop(temp_data_frame.index, axis=0,inplace=True)
#raise NotImplementedError()
```

<!-- grade_id: cell-9d04f770468db3c8, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert len(survey_df) == 63900
```

---

## Quay lại bước khám phá dữ liệu

### Với mỗi cột có kiểu dữ liệu không phải dạng số, các giá trị được phân bố như thế nào? (1đ)

Với các cột có kiểu dữ liệu không phải số, bạn sẽ tính:
- Tỉ lệ % (từ 0 đến 100) các giá trị thiếu 
- Số lượng các giá trị (các giá trị ở đây là các giá trị khác nhau và ta không xét giá trị thiếu): với cột mà ứng với câu hỏi dạng multichoice (ví dụ, cột "DevType"), mỗi giá trị có thể chứa nhiều choice (các choice được phân tách bởi dấu chấm phẩy), và việc đếm trực tiếp các giá trị không có nhiều ý nghĩa lắm vì số lượng tổ hợp các choice là khá nhiều; một cách khác tốt hơn mà bạn sẽ làm là đếm số lượng các choice
- Tỉ lệ % (từ 0 đến 100) của mỗi giá trị được sort theo tỉ lệ % giảm dần (ta không xét giá trị thiếu, tỉ lệ là tỉ lệ so với số lượng các giá trị không thiếu): bạn dùng dictionary để lưu, key là giá trị, value là tỉ lệ %; với cột mà ứng với câu hỏi dạng multichoice, cách làm tương tự như ở trên

Bạn sẽ lưu kết quả vào DataFrame `cate_col_info_df`, trong đó: 
- Tên của các cột là tên của các cột không phải số trong `survey_df`
- Tên của các dòng là: "missing_percentage", "num_values", "value_percentages"  

Để dễ nhìn, tất cả các giá trị bạn đều làm tròn với 1 chữ số thập phân bằng phương thức `.round(1)`.

Gợi ý: có thể bạn sẽ muốn dùng [phương thức `explode`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.Series.explode.html).

```python
# Các cột ứng với câu hỏi khảo sát multichoice
multichoice_cols = ['DevType', 'Gender', 'JobFactors', 
                    'LanguageWorkedWith', 'LanguageDesireNextYear', 
                    'MiscTechWorkedWith', 'MiscTechDesireNextYear', 
                    'NEWCollabToolsWorkedWith', 'NEWCollabToolsDesireNextYear', 
                    'PlatformWorkedWith', 'PlatformDesireNextYear', 
                    'NEWStuck']
```

<!-- grade_id: cell-c1cfa0a88177079d, points: 0, grade: false, solution: true -->
```python
pd.set_option('display.max_colwidth', 100) # Để dễ nhìn
pd.set_option('display.max_columns', None) # Để dễ nhìn
# YOUR CODE HERE
col_cate = survey_df.select_dtypes(include = None, exclude = ['float64'])
missing_percentage = round((((col_cate.isna().sum())/len(survey_df)) * 100),1)
missing_percentage = (missing_percentage.to_frame()).rename(columns={0:'missing_percentage'})
cate_col_info_df = missing_percentage.T

num_values_temp = {}
value_percent_temp = {}
for i in col_cate:
    value_percent_list = []
    if i in multichoice_cols:
        temp = survey_df[i].str.get_dummies(sep = ';').sum() #get_dummies này em tham khảo được trên stackoverflow,nó trả ra
        total_temp = temp.sum()                              # 1 series có index là tên dữ liệu và value là số lần xuất hiện
        num_values_list = list(temp.index)
        num_values_temp[i] = len(temp)
        value_percent_temp[i] = {}
        for j in num_values_list:
            value_percent_temp[i][j] = round((temp[j] / total_temp * 100),1)
    else:
        temp = survey_df[i].str.get_dummies().sum()
        total_temp = temp.sum()
        num_values_list = list(temp.index)
        num_values_temp[i] = len(temp)
        value_percent_temp[i] = {}
        for j in num_values_list:
            value_percent_temp[i][j] = round((temp[j] / total_temp * 100),1)      
num_values_final = pd.DataFrame.from_dict(num_values_temp,orient='index').rename(columns={0:'num_values'}).T
cate_col_info_df = cate_col_info_df.append(num_values_final)
value_percentage_final = []
for col, values in value_percent_temp.items():
    value_percentage_final.append(tuple((col, values)))
value_percentage_final = pd.DataFrame(value_percentage_final, index=value_percent_temp.keys(), columns=['0', 'value_percentages']).drop(columns=['0']).T
cate_col_info_df = cate_col_info_df.append(value_percentage_final)
#raise NotImplementedError()
cate_col_info_df
```

<!-- grade_id: cell-e748038de60d2ab7, points: 1, grade: true, solution: false -->
```python
# TEST
c = cate_col_info_df['MainBranch']
assert c.loc['missing_percentage'] == 0.5
assert c.loc['num_values'] == 5
assert c.loc['value_percentages']['I am a developer by profession'] == 73.5

c = cate_col_info_df['Hobbyist']
assert c.loc['missing_percentage'] == 0.1
assert c.loc['num_values'] == 2
assert c.loc['value_percentages']['Yes'] == 78.2

c = cate_col_info_df['DevType']
assert c.loc['missing_percentage'] == 23.6
assert c.loc['num_values'] == 23
assert c.loc['value_percentages']['Academic researcher'] == 2.2

c = cate_col_info_df['PlatformWorkedWith']
assert c.loc['missing_percentage'] == 16.5
assert c.loc['num_values'] == 16
assert c.loc['value_percentages']['Docker'] == 10.6
```

---

## Đặt câu hỏi 

Sau khi khám phá dữ liệu, ta đã hiểu hơn về dữ liệu. Bây giờ, ta hãy xem thử có câu hỏi nào có thể được trả lời bằng dữ liệu này.

**Một câu hỏi có thể có là:** Platform nào (Windows, Linux, Docker, AWS, ...) được yêu thích nhất, platform nào được yêu thích nhì, platform nào được yêu thích ba, ...?

Một platform được xem là được yêu thích nếu một người đã dùng platform này (cột "PlatformWorkedWith") và muốn tiếp tục dùng platform trong năm kế (cột "PlatformDesireNextYear").

**Trả lời được câu hỏi này sẽ** phần nào giúp ta định hướng là nên tập trung học platform nào để có thể chuẩn bị cho tương lai (mình nói "phần nào" vì ở đây dữ liệu chỉ giới hạn trong phạm vi những người làm khảo sát của StackOverflow).

---

## Tiền xử lý

Nếu bạn thấy cần thực hiện thêm thao tác tiền xử lý để chuẩn bị dữ liệu cho bước phân tích thì bạn làm ở đây. Bước này là không bắt buộc. 

```python
# YOUR CODE HERE (OPTION) 
```

---

## Phân tích dữ liệu (2.25đ)

Bây giờ, bạn sẽ thực hiện phân tích dữ liệu để trả lời cho câu hỏi ở trên. Cụ thể các bước như sau:
- Bước 1: tính Series `most_loved_platforms`, trong đó: 
    - Index là tên flatform (ở bước khám phá dữ liệu, bạn đã thấy có tất cả 16 platform)
    - Data là tỉ lệ % (từ 0 đến 100, được làm tròn với một chữ số thập phân bằng phương thức `round(1)`) được yêu thích (được sort giảm dần) 
- Bước 2: từ Series `most_loved_platforms`, bạn vẽ bar chart:
    - Bạn cho các bar nằm ngang (cho dễ nhìn)
    - Bạn đặt tên trục hoành là "Tỉ lệ %"

Code bước 1.

<!-- grade_id: cell-6d34c88e2fd46b33, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
platform = list(survey_df['PlatformDesireNextYear'].str.get_dummies(sep = ';').sum().index)
full_loved_platform = []
total_loved = 0
percent_platform = []
for i in range(0,16): #chứa thông tin các platform
    full_loved_platform.append([])
for i in range(0,16): # xử lý từng tên platform một
    full_loved_platform[i] = survey_df[survey_df['PlatformWorkedWith'].str.contains(platform[i]) & survey_df['PlatformDesireNextYear'].str.contains(platform[i])]
    total_loved += len(full_loved_platform[i])
for i in range(0,16): # tính phần trăm
    percent = round((len(full_loved_platform[i])/total_loved)*100, 1)
    percent_platform.append(percent)
most_loved_platforms = pd.Series(percent_platform, platform)
most_loved_platforms.sort_values(ascending=False, inplace=True)
#raise NotImplementedError()
most_loved_platforms
```

<!-- grade_id: cell-1d22896d0f3f2394, points: 2, grade: true, solution: false -->
```python
# TEST
assert len(most_loved_platforms) == 16
assert most_loved_platforms.loc['Linux'] == 20.2
assert most_loved_platforms.loc['Windows'] == 14.6
assert most_loved_platforms.loc['Docker'] == 12.3
```

Code bước 2.

<!-- grade_id: cell-5395c611d6065339, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
plt.figure(figsize=(10, 10))
plt.barh(most_loved_platforms.index, most_loved_platforms)
plt.xlabel('Tỉ lệ %')
#raise NotImplementedError()
```

Bạn đã hiểu tại sao mình khuyên bạn là nên tập làm quen dần với các câu lệnh của Linux chưa 😉 

---

## Đặt câu hỏi của bạn (1.5đ)

Bây giờ, đến lượt bạn phải tự suy nghĩ và đưa ra câu hỏi mà có thể trả lời bằng dữ liệu. Ngoài việc đưa ra câu hỏi, bạn cũng phải giải thích để người đọc thấy nếu trả lời được câu hỏi thì sẽ có lợi ích gì. Bạn nên sáng tạo một xíu, không nên đưa ra câu hỏi cùng dạng với câu hỏi của mình ở trên.  

<!-- grade_id: cell-2e7a92ae99ca7449, points: 1.5, grade: true, solution: true -->
Một vấn đề em nghĩ ra đó là: Tìm hiểu về mức lương của những người tham gia khảo sát và số năm làm việc chuyên nghiệp của họ. Liệu có phải trong lĩnh vực IT thì càng làm lâu, càng có nhiều kinh nghiệm làm việc, tiếp xúc với công nghệ lâu thì sẽ dễ có được mức thu nhập cao hơn so với những tân binh, những người trẻ mới vào nghề hay không?Hay việc thu nhập cao hay thấp còn có yếu tố nào ảnh hưởng tới nữa?

---

## Tiền xử lý để chuẩn bị dữ liệu cho bước phân tích để trả lời cho câu hỏi của bạn

Phần này là không bắt buộc. 

```python
# YOUR CODE HERE (OPTION) 
#bỏ đi những dòng mà YearsCodePro hoặc ConvertedComp trống
temp_df = survey_df[survey_df['YearsCodePro'].isnull() | survey_df['ConvertedComp'].isnull()]
survey_df.drop(temp_df.index,axis=0,inplace=True)
```

---

## Phân tích dữ liệu để ra câu trả lời cho câu hỏi của bạn (2đ)

<!-- grade_id: cell-83e11892184cf811, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
df_to_draw = survey_df[['ConvertedComp','YearsCodePro']]
m, b = np.polyfit(df_to_draw['YearsCodePro'], df_to_draw['ConvertedComp'],1)
plt.figure(figsize=(30,30))
plt.scatter(df_to_draw['YearsCodePro'],df_to_draw['ConvertedComp'], s = 15, c ='black')
plt.xlabel('Số năm kinh nghiệm')
plt.ylabel('Mức lương')
plt.plot(df_to_draw['YearsCodePro'], m*df_to_draw['YearsCodePro'] + b)
#raise NotImplementedError()
#nhận xét về câu hỏi đặt ra: dựa vào biểu đồ scatter ở dưới thì có thể thấy được rằng không phải luôn luôn nhất thiết tuổi nghề 
#phải nhiều, phải có nhiều kinh nghiệm thì thu nhập mới cao, vẫn có khá nhiều trường hợp tuổi nghề còn trẻ nhưng mức thu nhập
#vẫn khá cao so với con số trung bình. Tuy nhiên nhìn vào đường kẻ hồi quy thì ta cũng thấy được rằng đúng là nếu càng làm lâu
#trong nghề thì mức thu nhập nói chung vẫn sẽ tăng nên nếu kết luận là tuổi nghề cao thì mức lương cao cũng mang tính tổng quát
#nhất định dựa trên số liệu(do có một số outlier nên biểu đồ trông hơi khó để nhìn ra)
```