# Đồ án cuối kỳ - Môn Lập trình khoa học dữ liệu
# Giáo viên giảng dạy: Thầy Trần Trung Kiên
## Sinh viên thực hiện:

                        Tạ Võ Anh Khuê - 19120551

                        Nguyễn Minh Long - 19120568

---

## Import

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd 
```

---

## Thu thập dữ liệu

Dữ liệu được sử dụng trong bài tập này là dữ liệu về các phàn nàn, khiếu nại của người dùng được gửi tới CFBP(Consumer Financial Protection Bureau) - tạm gọi là Cục Bảo Vệ Tài Chính Cho Người Tiêu Dùng. Dữ liệu được lấy [ở đây](https://www.kaggle.com/meetnagadia/consumer-complaint-finance). 

Cục Bảo Vệ Tài Chính Cho Người Tiêu Dùng (CFPB) là một cơ quan của chính phủ liên bang Mỹ được thành lập để bảo đảm sự hiệu quả của những sản phẩm và dịch vụ tài chính cho mọi người như mua nhà, chọn lựa một thẻ tín dụng, gửi tiền cho người thân ở nước ngoài, hoặc khi sử dụng bất cứ loại sản phẩm tài chánh khác ở Mỹ.
Cục có chức năng giám sát hoạt động của các ngân hàng, liên hiệp tín dụng (credit union), những công ty tài chính khác và thi hành những quy luật tài chính để bảo vệ người dùng tránh khỏi các hành vi bất công, dối trá và lạm dụng.

Đôi nét về dữ liệu này: Dữ liệu thể hiện các complaint(khiếu nại, phàn nàn) của người dùng về các vấn đề tài chính mà họ gặp phải. Người dùng trình báo lên Cục Bảo Vệ(CFBP) và cục có trách nhiệm liên lạc và trao đổi, chuyển tiếp thông tin giúp người dùng tới các công ty liên quan. Dữ liệu được công khai ghi nhận(published) bởi CFPB sau khi mà các công ty phản hồi, xác nhận có liên quan đến các vấn đề của người dùng hoặc sau 15 ngày, tùy điều kiện nào đến trước.

---

## Khám phá dữ liệu - Currently Working

## Đọc dữ liệu từ file

Dữ liệu được đặt chung cấp folder của file noteboook, đọc vào bằng hàm read_csv đơn giản, không có gì phức tạp trong phần này.

<!-- grade_id: cell-d46e97079096473c, points: 0.25, grade: true, solution: true -->
```python
df = pd.read_csv('complaints.csv')
df.head()
```

Output:
```
D:\Miniconda\envs\min_ds-env\lib\site-packages\IPython\core\interactiveshell.py:3172: DtypeWarning: Columns (9) have mixed types.Specify dtype option on import or set low_memory=False.
  has_raised = await self.run_ast_nodes(code_ast.body, cell_name,
  Date received                                            Product  \
0    2019-06-13  Credit reporting, credit repair services, or o...   
1    2019-11-01                              Vehicle loan or lease   
2    2019-04-01  Credit reporting, credit repair services, or o...   
3    2019-04-10                        Credit card or prepaid card   
4    2020-04-21  Credit reporting, credit repair services, or o...   

         Sub-product                                 Issue  \
0   Credit reporting  Incorrect information on your report   
1               Loan           Struggling to pay your loan   
2   Credit reporting  Incorrect information on your report   
3  Store credit card                  Closing your account   
4   Credit reporti
```

Xử lý một chút về cột dữ liệu Tags, ta có thể thấy giá trị cột Tags bị thiếu quá nhiều, gần 90% giá trị bị thiếu ở cột này => Ta có thể bỏ cột này đi vì gần như nó không đóng góp gì cho việc phân tích dữ liệu sau này

```python
round(df['Tags'].isna().sum()/len(df) * 100, 2)
```

Output:
```
87.88
```

```python
df = df.drop('Tags',axis=1)
```

Xem lại dữ liệu 1 lần nữa sau khi đã bỏ cột Tags

```python
df.head()
```

Output:
```
Date received                                            Product  \
0    2019-06-13  Credit reporting, credit repair services, or o...   
1    2019-11-01                              Vehicle loan or lease   
2    2019-04-01  Credit reporting, credit repair services, or o...   
3    2019-04-10                        Credit card or prepaid card   
4    2020-04-21  Credit reporting, credit repair services, or o...   

         Sub-product                                 Issue  \
0   Credit reporting  Incorrect information on your report   
1               Loan           Struggling to pay your loan   
2   Credit reporting  Incorrect information on your report   
3  Store credit card                  Closing your account   
4   Credit reporting  Incorrect information on your report   

                             Sub-issue  \
0  Information belongs to someone else   
1     Denied request to lower payments   
2             Account status incorrect   
3             Can't close your account  
```

### Mỗi cột dữ liệu có ý nghĩa là gì?

Data có tổng cộng 17 cột, với các ý nghĩa cơ bản của từng cột như sau:
* Date received: ngày nhận được khiếu nại
* Product: loại sản phẩm, dịch vụ khiếu nại,....
* Sub product: thêm thông tin, làm rõ hơn về sản phẩm, dịch vụ,... trong product
* Issue: vấn đề người dùng gặp phải
* Sub-issue: thông tin thêm, làm rõ về vấn đề gặp phải trong issue
* Consumer complaint narrative: miêu tả, phàn nàn, khiếu nại của người dùng đối với vấn đề họ gặp phải
* Company public response: phản hồi chính thức, công khai của phía công ty
* Company: tên công ty
* State: tên bang
* ZIP code : mã zip ( giống mã bưu điện )
* Consumer consent provided: người dùng có đồng ý public(công khai) vấn đề của mình ra hay không
* Submitted via: khiếu nại được gửi thông qua dạng nào(web,fax,email,...)
* Date sent to company: ngày khiếu nại được gửi đến công ty
* Company response to consumer: tình trạng phản hồi, giải quyết của công ty đối với người khiếu nại
* Timely response: trả lời đúng thời hạn
* Consumer disputed?: người khiếu nại có mở tranh chấp không(dạng đòi bồi thường này kia)
* Complaint ID: id của đơn khiếu nại

### Dữ liệu có bao nhiêu dòng và bao nhiêu cột

<!-- grade_id: cell-8338f3610f6c5002, points: 0, grade: false, solution: true -->
```python
size = df.shape
size
```

Output:
```
(2388981, 17)
```

### Xem các dòng có bị lặp không thông qua giá trị của cột 'Complaint ID'

```python
num_duplicated_rows = df.duplicated(subset=['Complaint ID']).sum()
num_duplicated_rows
```

Output:
```
0
```

#### Tiếp theo kiểm tra thử cột Consumer disputed? , ta thấy có các giá trị như sau

```python
df['Consumer disputed?'].unique()
```

Output:
```
array([nan, 'No', 'Yes'], dtype=object)
```

##### Trong trường hợp này ta có thể hiểu được rằng NaN sẽ có ý nghĩa tương đương với No, vì tâm lý người dùng khi họ không còn vấn đề gì cần đòi hỏi tranh chấp, giải quyết với công ty nữa thì họ sẽ điền là 'No', nhưng có thể một số người họ quên hoặc không chú ý khi điền thông tin này nên họ để trống dẫn đến NaN. Vậy nên ta có thể thay các giá trị NaN là No luôn

```python
df['Consumer disputed?'] = df['Consumer disputed?'].fillna('No')
df['Consumer disputed?'].unique()
```

### Mỗi cột có kiểu dữ liệu gì?

```python
df.dtypes
```

Output:
```
Date received                   object
Product                         object
Sub-product                     object
Issue                           object
Sub-issue                       object
Consumer complaint narrative    object
Company public response         object
Company                         object
State                           object
ZIP code                        object
Consumer consent provided?      object
Submitted via                   object
Date sent to company            object
Company response to consumer    object
Timely response?                object
Consumer disputed?              object
Complaint ID                     int64
dtype: object
```

Xem kiểu dữ liệu thật sự bên trong mỗi cột.

```python
def open_object_dtype(s):
    dtypes = set()
    a = s.apply(lambda x: type(x))
    dtypes = set((a.unique()))
    return dtypes   
col = df.columns
for i in range(len(col)):
    temp = open_object_dtype(df[col[i]])
    print(str(col[i])+' ' +str(temp))
```

Output:
```
Date received {<class 'str'>}
Product {<class 'str'>}
Sub-product {<class 'str'>, <class 'float'>}
Issue {<class 'str'>}
Sub-issue {<class 'str'>, <class 'float'>}
Consumer complaint narrative {<class 'str'>, <class 'float'>}
Company public response {<class 'str'>, <class 'float'>}
Company {<class 'str'>}
State {<class 'str'>, <class 'float'>}
ZIP code {<class 'str'>, <class 'float'>}
Consumer consent provided? {<class 'str'>, <class 'float'>}
Submitted via {<class 'str'>}
Date sent to company {<class 'str'>}
Company response to consumer {<class 'str'>, <class 'float'>}
Timely response? {<class 'str'>}
Consumer disputed? {<class 'str'>, <class 'float'>}
Complaint ID {<class 'int'>}
```

### Chuyển 2 cột 'Date received' và 'Date sent to company' sang dạng Datetime

```python
df['Date received'] = pd.to_datetime(df['Date received'],format='%Y-%m-%d')
df['Date sent to company'] = pd.to_datetime(df['Date sent to company'],format='%Y-%m-%d')
```

```python
df.dtypes
```

Output:
```
Date received                   datetime64[ns]
Product                                 object
Sub-product                             object
Issue                                   object
Sub-issue                               object
Consumer complaint narrative            object
Company public response                 object
Company                                 object
State                                   object
ZIP code                                object
Consumer consent provided?              object
Submitted via                           object
Date sent to company            datetime64[ns]
Company response to consumer            object
Timely response?                        object
Consumer disputed?                      object
Complaint ID                             int64
dtype: object
```

#### Xét giá trị trong cột 'Date received' có lớn hơn giá trị trong cột 'Date sent to company'. Bởi vì nếu ngày received mà xảy ra sau ngày sent to company là vô lý. Ta sẽ xóa các dòng vô lý đó.

```python
df = df.loc[df['Date received'] <= df['Date sent to company']]
df.shape
```

Output:
```
(2388981, 17)
```

```python
df['Consumer consent provided?'].value_counts(dropna=False)
```

Output:
```
Consent provided        835565
Consent not provided    759039
NaN                     726949
Other                    63839
Consent withdrawn         3589
Name: Consumer consent provided?, dtype: int64
```

---

## Đặt câu hỏi 

Sau khi khám phá dữ liệu, ta đã hiểu hơn về dữ liệu. Bây giờ, ta hãy xem thử có câu hỏi nào có thể được trả lời bằng dữ liệu này.

**Một câu hỏi có thể có là:** Platform nào (Windows, Linux, Docker, AWS, ...) được yêu thích nhất, platform nào được yêu thích nhì, platform nào được yêu thích ba, ...?

Một platform được xem là được yêu thích nếu một người đã dùng platform này (cột "PlatformWorkedWith") và muốn tiếp tục dùng platform trong năm kế (cột "PlatformDesireNextYear").

**Trả lời được câu hỏi này sẽ** phần nào giúp ta định hướng là nên tập trung học platform nào để có thể chuẩn bị cho tương lai (mình nói "phần nào" vì ở đây dữ liệu chỉ giới hạn trong phạm vi những người làm khảo sát của StackOverflow).

---

## Tiền xử lý

Nếu bạn thấy cần thực hiện thêm thao tác tiền xử lý để chuẩn bị dữ liệu cho bước phân tích thì bạn làm ở đây. Bước này là không bắt buộc. 

```python
# YOUR CODE HERE (OPTION) 
```

---

## Phân tích dữ liệu (2.25đ)

Bây giờ, bạn sẽ thực hiện phân tích dữ liệu để trả lời cho câu hỏi ở trên. Cụ thể các bước như sau:
- Bước 1: tính Series `most_loved_platforms`, trong đó: 
    - Index là tên flatform (ở bước khám phá dữ liệu, bạn đã thấy có tất cả 16 platform)
    - Data là tỉ lệ % (từ 0 đến 100, được làm tròn với một chữ số thập phân bằng phương thức `round(1)`) được yêu thích (được sort giảm dần) 
- Bước 2: từ Series `most_loved_platforms`, bạn vẽ bar chart:
    - Bạn cho các bar nằm ngang (cho dễ nhìn)
    - Bạn đặt tên trục hoành là "Tỉ lệ %"

Code bước 1.

<!-- grade_id: cell-6d34c88e2fd46b33, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
platform = list(survey_df['PlatformDesireNextYear'].str.get_dummies(sep = ';').sum().index)
full_loved_platform = []
total_loved = 0
percent_platform = []
for i in range(0,16): #chứa thông tin các platform
    full_loved_platform.append([])
for i in range(0,16): # xử lý từng tên platform một
    full_loved_platform[i] = survey_df[survey_df['PlatformWorkedWith'].str.contains(platform[i]) & survey_df['PlatformDesireNextYear'].str.contains(platform[i])]
    total_loved += len(full_loved_platform[i])
for i in range(0,16): # tính phần trăm
    percent = round((len(full_loved_platform[i])/total_loved)*100, 1)
    percent_platform.append(percent)
most_loved_platforms = pd.Series(percent_platform, platform)
most_loved_platforms.sort_values(ascending=False, inplace=True)
#raise NotImplementedError()
most_loved_platforms
```

<!-- grade_id: cell-1d22896d0f3f2394, points: 2, grade: true, solution: false -->
```python
# TEST
assert len(most_loved_platforms) == 16
assert most_loved_platforms.loc['Linux'] == 20.2
assert most_loved_platforms.loc['Windows'] == 14.6
assert most_loved_platforms.loc['Docker'] == 12.3
```

Code bước 2.

<!-- grade_id: cell-5395c611d6065339, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
plt.figure(figsize=(10, 10))
plt.barh(most_loved_platforms.index, most_loved_platforms)
plt.xlabel('Tỉ lệ %')
#raise NotImplementedError()
```

Bạn đã hiểu tại sao mình khuyên bạn là nên tập làm quen dần với các câu lệnh của Linux chưa 😉 

---

## Đặt câu hỏi của bạn (1.5đ)

Bây giờ, đến lượt bạn phải tự suy nghĩ và đưa ra câu hỏi mà có thể trả lời bằng dữ liệu. Ngoài việc đưa ra câu hỏi, bạn cũng phải giải thích để người đọc thấy nếu trả lời được câu hỏi thì sẽ có lợi ích gì. Bạn nên sáng tạo một xíu, không nên đưa ra câu hỏi cùng dạng với câu hỏi của mình ở trên.  

<!-- grade_id: cell-2e7a92ae99ca7449, points: 1.5, grade: true, solution: true -->
Một vấn đề em nghĩ ra đó là: Tìm hiểu về mức lương của những người tham gia khảo sát và số năm làm việc chuyên nghiệp của họ. Liệu có phải trong lĩnh vực IT thì càng làm lâu, càng có nhiều kinh nghiệm làm việc, tiếp xúc với công nghệ lâu thì sẽ dễ có được mức thu nhập cao hơn so với những tân binh, những người trẻ mới vào nghề hay không?Hay việc thu nhập cao hay thấp còn có yếu tố nào ảnh hưởng tới nữa?

---

## Tiền xử lý để chuẩn bị dữ liệu cho bước phân tích để trả lời cho câu hỏi của bạn

Phần này là không bắt buộc. 

```python
# YOUR CODE HERE (OPTION) 
#bỏ đi những dòng mà YearsCodePro hoặc ConvertedComp trống
temp_df = survey_df[survey_df['YearsCodePro'].isnull() | survey_df['ConvertedComp'].isnull()]
survey_df.drop(temp_df.index,axis=0,inplace=True)
```

---

## Phân tích dữ liệu để ra câu trả lời cho câu hỏi của bạn (2đ)

<!-- grade_id: cell-83e11892184cf811, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
df_to_draw = survey_df[['ConvertedComp','YearsCodePro']]
m, b = np.polyfit(df_to_draw['YearsCodePro'], df_to_draw['ConvertedComp'],1)
plt.figure(figsize=(30,30))
plt.scatter(df_to_draw['YearsCodePro'],df_to_draw['ConvertedComp'], s = 15, c ='black')
plt.xlabel('Số năm kinh nghiệm')
plt.ylabel('Mức lương')
plt.plot(df_to_draw['YearsCodePro'], m*df_to_draw['YearsCodePro'] + b)
#raise NotImplementedError()
#nhận xét về câu hỏi đặt ra: dựa vào biểu đồ scatter ở dưới thì có thể thấy được rằng không phải luôn luôn nhất thiết tuổi nghề 
#phải nhiều, phải có nhiều kinh nghiệm thì thu nhập mới cao, vẫn có khá nhiều trường hợp tuổi nghề còn trẻ nhưng mức thu nhập
#vẫn khá cao so với con số trung bình. Tuy nhiên nhìn vào đường kẻ hồi quy thì ta cũng thấy được rằng đúng là nếu càng làm lâu
#trong nghề thì mức thu nhập nói chung vẫn sẽ tăng nên nếu kết luận là tuổi nghề cao thì mức lương cao cũng mang tính tổng quát
#nhất định dựa trên số liệu(do có một số outlier nên biểu đồ trông hơi khó để nhìn ra)
```